"""
City Information Configuration Module

Manages city data, measurements, pricing, and user interaction
utilities for the city information system.
"""
# I AM NOT DONE

cityName = None  # TODO: Set the city name

length = None   # TODO: Set a length value for city planning
width = None    # TODO: Set a width value
area = None     # TODO: Calculate length * width

temperature = None  # TODO: Set to 98.6


def greetUser():
    """Print a welcome message to the user."""
    raise NotImplementedError()


def addNumbers(a, b):
    """Return the sum of two numbers."""
    raise NotImplementedError()


total = None  # TODO: Call addNumbers(15, 27) and store the result


def squareNumber(number):
    """Return the number multiplied by itself."""
    raise NotImplementedError()


price = None       # TODO: Set to 49.99
quantity = None    # TODO: Set to 3
orderCost = None   # TODO: Calculate price * quantity


def convertToMeters(feet):
    """Convert feet to meters (1 foot = 0.3048 meters)."""
    raise NotImplementedError()


def getFullName(firstName, lastName):
    """Return first and last name combined with a space."""
    raise NotImplementedError()


score = None  # TODO: Set to 0


def addPoints(number):
    """Return score plus the given number."""
    raise NotImplementedError()
