"""
Text Formatting and Analysis Module

String manipulation, formatting, pattern matching,
and text processing utilities.
"""
# I AM NOT DONE

itemName = "croissant"
quantity = 12
price = 3.5

# TODO: Use an f-string to format: "Sold 12 croissants for $42.00"
formattedSale = None

# TODO: Use .join() to combine ["a", "b", "c"] with " - "
joined = None

# TODO: Use .split() to break "apple,banana,grape" into a list
fruits = "apple,banana,grape"
splitFruits = None

# TODO: Use .strip() to remove whitespace from "  hello world  "
cleaned = None

# TODO: Use .replace() to swap "cat" with "dog" in "the cat sat"
replaced = None

# TODO: Use .find() to get the index of "world" in "hello world"
findIndex = None

# TODO: Use .isalpha(), .isdigit(), .isalnum() on "hello123"
alpha = None   # .isalpha()
digit = None   # .isdigit()
alnum = None   # .isalnum()

# TODO: Use an f-string with formatting: format 1/3 as percentage with 1 decimal
fraction = 1/3
asPercent = None  # Should be "33.3%"

# TODO: Format 1234567 with commas
bigNumber = 1234567
withCommas = None  # Should be "1,234,567"

import re

# TODO: Use re.findall to find all words starting with "b"
text = "the baker bakes beautiful bagels"
bWords = None
