"""
Data Transformation Service

Module for transforming, filtering, and processing
collections of data using functional patterns.
"""
# I AM NOT DONE

numbers = None  # TODO: Set to list(range(1, 11)) — [1, 2, 3, ..., 10]

# TODO: Use a list comprehension to square each number
squares = None  # [1, 4, 9, ..., 100]

# TODO: Use a list comprehension to get only even numbers from numbers
evens = None  # [2, 4, 6, 8, 10]

words = None  # TODO: Set to at least 5 words, e.g. ["apple", "banana", ...]

# TODO: Use a list comprehension to get the length of each word
wordLengths = None

# TODO: Use a list comprehension to get only words that start with "a"
aWords = None

# TODO: Use a dict comprehension to map each word -> its length
wordLengthDict = None

# TODO: Use a set comprehension to get unique first letters from words
firstLetters = None

# TODO: Use a list comprehension with enumerate to create
# ["0: croissant", "1: muffin", ...] from items list
items = ["croissant", "muffin", "bagel", "cookie"]
numberedItems = None
