"""
Mathematical Computation Module

Utility functions and constants for mathematical
operations, statistics, and random generation.
"""
# I AM NOT DONE

import math
import random
import statistics

# TODO: Use math.ceil() and math.floor()
ceilResult = None   # math.ceil(3.2)
floorResult = None  # math.floor(3.8)

# TODO: Use math.sqrt() and math.pow()
sqrtResult = None   # math.sqrt(144)
powResult = None    # math.pow(2, 10)

# TODO: Use math.pi to calculate area of a circle with radius 3
# Formula: pi * radius ** 2
circleArea = None

# TODO: Use random.randint() to generate a random int between 1 and 100
randomInt = None

# TODO: Use random.choice() to pick from ["red", "blue", "green", "yellow"]
randomColor = None

# TODO: Use random.sample() to pick 2 unique items from [10, 20, 30, 40, 50]
sample = None

# TODO: Use statistics.mean() and statistics.median() on [10, 20, 30, 40, 50]
data = [10, 20, 30, 40, 50]
meanValue = None
medianValue = None

# TODO: Use round() to round 3.14159 to 2 decimal places
rounded = None

# TODO: Use min(), max(), and sum() on [4, 7, 2, 9, 5]
nums = [4, 7, 2, 9, 5]
minNum = None
maxNum = None
sumNums = None

# TODO: Use math.factorial() on 5
factorialResult = None

# TODO: Use math.gcd() on 12 and 8
gcdResult = None
