"""
Date and Time Management Module

Utilities for date parsing, formatting, scheduling,
and time-based calculations.
"""
# I AM NOT DONE

from datetime import datetime, date, timedelta

# TODO: Get the current date and time
now = None

# TODO: Extract year, month, day from now
currentYear = None
currentMonth = None
currentDay = None

# TODO: Create a specific date: Christmas 2026
christmas = None

# TODO: Format current date as "June 04, 2026"
formattedDate = None

# TODO: Parse "2026-12-25" into a datetime object
parsedDate = None

# TODO: Calculate the number of days until Christmas
daysUntilChristmas = None

# TODO: Create a date 7 days from now
oneWeekLater = None

# TODO: Create a date 30 days ago
thirtyDaysAgo = None

# TODO: Use strftime to format as "Thursday, June 04, 2026"
todayFormatted = None

# TODO: Compare dates: is christmas after today?
isChristmasFuture = None

# TODO: Create a specific time (e.g., 2:30 PM)
from datetime import time
specificTime = None

# TODO: Calculate the difference between two dates in days
date1 = date(2026, 1, 1)
date2 = date(2026, 12, 31)
daysBetween = None
