"""
User Profile Management Module

Handles user profile data types, display utilities,
and common profile calculations.
"""
# I AM NOT DONE

myString = None  # TODO: Set to a string value
myInt = None     # TODO: Set to an integer value
myFloat = None   # TODO: Set to a float value
myBool = None    # TODO: Set to a boolean value
myList = None    # TODO: Set to a list value

favoriteFoods = None  # TODO: Set to a list with at least 4 items

isLoggedIn = None  # TODO: Set to False, then True on the next line


def describeUser(userName, age):
    """Print a sentence using the user's name and age."""
    raise NotImplementedError()


def calcTip(billTotal, tipPercent):
    """Return the tip amount for a given bill total and tip percentage."""
    raise NotImplementedError()


weekdays = None  # TODO: Set to a list of all 5 weekdays


def isAdult(age):
    """Return True if age is 18 or older, False otherwise."""
    raise NotImplementedError()


gpa = None             # TODO: Set to a float value
graduationYear = None  # TODO: Set to an integer value
