"""
Game State Management Module

Manages global game state including player stats,
scoring, inventory, and game control functions.
"""
# I AM NOT DONE

appName = None  # TODO: Set to the application name


def localTest():
    """Create a local variable and return it."""
    raise NotImplementedError()


lives = None  # TODO: Set to 3


def loseLife():
    """Decrement lives by 1 using the global keyword."""
    raise NotImplementedError()


playerName = None  # TODO: Set to "Player One"


def changeName():
    """Change playerName using the global keyword."""
    raise NotImplementedError()


totalSales = None  # TODO: Set to 0


def makeSale(amount):
    """Add amount to totalSales using the global keyword."""
    raise NotImplementedError()


color = None  # TODO: Set to "blue"


def localColor():
    """Return a local 'color' set to 'red' without changing the global."""
    raise NotImplementedError()


counter = None  # TODO: Set to 0


def incrementCounter():
    """Add 1 to counter using the global keyword."""
    raise NotImplementedError()


def getDiscount():
    """Return the local variable discountRate set to 0.15."""
    raise NotImplementedError()


isGameOver = None  # TODO: Set to False


def endGame():
    """Set isGameOver to True using the global keyword."""
    raise NotImplementedError()


def outerFunction():
    """
    Create a nested function that returns a variable from the outer scope.
    Return the result of calling the inner function.
    """
    raise NotImplementedError()


budget = None  # TODO: Set to 500


def spendMoney(cost):
    """Subtract cost from budget using the global keyword."""
    raise NotImplementedError()
