"""
Data Processing Module

Collection of data processing utilities for working with
lists, loops, and data transformations.
"""
# I AM NOT DONE

animals = None        # TODO: Set to a list of 5 animals
printedAnimals = []   # TODO: Use a for loop to copy animals into here

oneToTen = []    # TODO: Use a for loop with range() to fill with 1-10
countUp = []     # TODO: Use a while loop to fill with 1-5

prices = None      # TODO: Set to a list of at least 4 floats
priceTotal = 0.0   # TODO: Use a for loop to sum the prices

evens = []       # TODO: Use a loop to fill with even numbers 0-20
countDown = []   # TODO: Start at 100, subtract 15 until below 0

names = None         # TODO: Set to a list with at least 5 names, some starting with "A"
filteredNames = []   # TODO: Loop with continue to exclude names starting with "A"

numbers = None     # TODO: Set to a list that includes a number > 50
foundNumber = None # TODO: Loop with break to find the first number > 50

multiplicationTable = []  # TODO: Loop to fill with 7x1 through 7x10

attempts = 0
lockMessage = None  # TODO: While loop, at 3 attempts set this to "Locked out!"

mixedItems = None  # TODO: Set to a list where some values are None
realItems = []     # TODO: Loop with continue to keep only non-None values

eCount = 0  # TODO: Count how many times "e" appears in "intermediate"
