"""
Object-Oriented Models Module

Defines domain models for a vehicle management and
pet adoption system using classes and inheritance.
"""
# I AM NOT DONE

# TODO: Create a Car class with make, model, year and a describe() method.
#       describe() should return "[make] [model]".

car1 = None  # TODO: Create a Car instance
car2 = None  # TODO: Create another Car instance with a different make

# TODO: Create an Animal class with name, sound, and a speak() method.
#       speak() should return "[name] says [sound]".

# TODO: Create a Pet class that inherits from Animal, adds ownerName,
#       and adds greetOwner() returning "[name] runs to [ownerName]!".

myPet = None  # TODO: Create a Pet instance

# TODO: Create a BankAccount class with accountHolder, balance,
#       deposit(amount), and withdraw(amount) methods.
#       withdraw returns "Insufficient funds." if amount > balance.

# TODO: Create a SavingsAccount class that inherits from BankAccount,
#       adds interestRate, and adds applyInterest().

myAccount = None  # TODO: Create a BankAccount instance
# TODO: Change accountHolder directly on myAccount

animalOne = None  # TODO: Create an Animal instance
animalTwo = None  # TODO: Create another Animal instance
# TODO: Change animalOne's name directly
