"""
Text and Data Transformation Utilities

Utility functions for string manipulation, list operations,
dictionary management, and data type conversions.
"""
# I AM NOT DONE

paddedString = None   # TODO: Set to a string with leading/trailing spaces
strippedString = None # TODO: Use strip() on paddedString

lowerTitle = None  # TODO: Set to a lowercase string
titledString = None # TODO: Use title() on lowerTitle

original = None  # TODO: Set to a string
replaced = None  # TODO: Use replace() to swap a word

sentence = None  # TODO: Set to a multi-word sentence
wordList = None  # TODO: Use split() on sentence

unsortedNums = None  # TODO: Set to a list of at least 5 numbers out of order
sortedNums = None    # TODO: Sort the list

editList = None  # TODO: Set to a list with at least 5 items
# TODO: Use remove() to delete an item, then append() to add a new one

personDict = None  # TODO: Create a dict with at least 3 key-value pairs
dictKeys = None    # TODO: Extract keys from personDict
dictValues = None  # TODO: Extract values from personDict

missingValue = None  # TODO: Use get() on personDict for a missing key "nickname"

updatedDict = None  # TODO: Use update() then pop() on personDict

numString = None  # TODO: Set to a string containing a number like "77"
doubledNum = None # TODO: Convert to int and multiply by 2

dupList = None    # TODO: Set to a list with duplicate values
uniqueList = None # TODO: Use set() to remove duplicates

wordBank = None  # TODO: Set to a list of at least 6 words, some with "a"
aCount = 0       # TODO: Count how many words contain the letter "a"
